# CS6140 Assignments

**Instructions**
1. In each assignment cell, look for the block:
 ```
  #BEGIN YOUR CODE
  raise NotImplementedError.new()
  #END YOUR CODE
 ```
1. Replace this block with your solution.
1. Test your solution by running the cells following your block (indicated by ##TEST##)
1. Click the "Validate" button above to validate the work.

**Notes**
* You may add other cells and functions as needed
* Keep all code in the same notebook
* In order to receive credit, code must "Validate" on the JupyterHub server

---

```ruby

```

# Assignment 2: Clustering

```ruby
require './assignment_lib'
```

Output:
```
false
```

### Note on data format

A "Dataset” is a hash with in which ```"data"``` is an array of examples. For example:

```ruby
dataset = {"data" => examples, "features" => ["x1", "x2"]}
```

The examples are represented as an array of hash where each example is a hash in which ```"features"``` is a hash of features and “label” is the true label. For example:

```ruby
{"features"=>{"x1"=>-3.925613761029373, "x2"=>4.567177634742035}, "label"=>0, "cluster"=>0},
```

## Question 1.1: K-Means (25 points)

The following dataset has 5 clusters. Generate a dataset with 5 means in two dimensions, with means $\mu$ and variance $\sigma$ defined as follows:
* $\mu_1 = (-4,4)$, $\sigma_1 = 1$
* $\mu_2 = (-4,-4)$, $\sigma_2 = 1$
* $\mu_3 = (4,4)$, $\sigma_3 = 1$
* $\mu_4 = (4,-4)$, $\sigma_4 = 1$
* $\mu_5 = (0,0)$, $\sigma_5 = 1$

Name your features ```x1``` and ```x2```.

Generate 100 points in each cluster and plot each cluster with a different color.  Because we are generating the dataset, we know what the true clustering should be and we set this as the ```label``` field. We will initiatlize the ```cluster``` field to the same value for now, but after each iteration, we will reassign ```cluster```. 

#### Generating an example from a normal distribution
To generate a value from a normal distribution, use the following:

```ruby
r = Distribution::Normal.rng(mean,stdev)
x = r.call()
```

where ```mean``` and ```stdev``` are ```Float```s representing the mean and the standard deviation of a single-variable normal distribution. If we wanted to generate an example from mean $\mu_1$, then we could write the following:

```ruby
r_11 = Distribution::Normal.rng(-4.0,1.0)
r_12 = Distribution::Normal.rng(4.0,1.0)

x1 = r_11.call()
x2 = r_12.call()
example = [x1, x2]
```
which results in ```example``` being ```[-4.64951882400697, 4.405135059361256]```.


#### Generating the cluster dataset
The result of the ```create_cluster_dataset``` function returns an array of examples, as follows:

```ruby
dataset = create_cluster_dataset()
examples = dataset["data"]
examples[0,2]
```

The first two examples are as follows
```
[
    {"features"=>{"x1"=>-3.925613761029373, "x2"=>4.567177634742035}, "label"=>0, "cluster"=>0}, 
    {"features"=>{"x1"=>-4.796103206668446, "x2"=>3.8075466648390637}, "label"=>0, "cluster"=>0}
]
```

<!-- grade_id: cell-1cb0100f7efe89e6, points: 0, grade: false, solution: true -->
```ruby
def create_cluster_dataset()
  n = 100
  clusters = [[-4,4], [-4,-4], [4,4], [4,-4], [0,0]]
  examples = []
  # BEGIN YOUR CODE
  j = 0
  $num = clusters.length()
  while  j < $num  do
        a = clusters[j][0]
        b = clusters[j][1]
        $i = 0           
        while  $i < n do          
            r_11 = Distribution::Normal.rng(a,1.0)
            r_12 = Distribution::Normal.rng(b,1.0)
            x1 = r_11.call()
            x2 = r_12.call()
            examples << {"features"=>{"x1"=>x1, "x2"=>x2}, "label"=>j, "cluster"=>j}
            $i +=1
        end 
        j +=1
  end 
 
  
  #END YOUR CODE
  
  dataset = {"data" => examples, "features" => ["x1", "x2"]}
  return dataset
end
```

Output:
```
:create_cluster_dataset
```

<!-- grade_id: cell-1ac6e84306bf8a37, points: 5, grade: true, solution: false -->
```ruby
def test_11_1()
  dataset = create_cluster_dataset()
  assert_false(dataset["data"].empty?)
  assert_equal 500, dataset["data"].size
  plot_clusters(dataset["data"])
end
test_11_1()
```

Output:
```
#<Nyaplot::Plot:0x0000555f5698f750 @properties={:diagrams=>[#<Nyaplot::Diagram:0x0000555f56152ab0 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>0, :color=>"rgb(27,158,119)", :tooltip_contents=>[0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]}, :data=>"14d7bef5-d907-4fc9-bb0c-4096de658076"}, @xrange=[-7.256648868291618, -2.062138340319478], @yrange=[1.5467470070381504, 6.006961662001062]>, #<Nyaplot::Diagram:0x0000555f57e2c230 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>1, :color=>"rgb(217,95,2)", :tooltip_contents=>[1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1,
```

The test below verifies that each cluster has 100 examples, both by looking at the ```cluster``` and ```label``` field. 

<!-- grade_id: cell-a6100841ca0c2971, points: 10, grade: true, solution: false -->
```ruby
def test_11_2()
  dataset = create_cluster_dataset()
  examples = dataset["data"]
  counts = examples.group_by {|x| x["cluster"]}
  counts.each_key {|k| counts[k] = counts[k].size}
  assert_equal 5, counts.size
  assert counts.values.all? {|v| v == 100}
  
  counts = examples.group_by {|x| x["label"]}
  counts.each_key {|k| counts[k] = counts[k].size}
  assert_equal 5, counts.size
  assert counts.values.all? {|v| v == 100}

end
test_11_2()
```

The test below verifies that the means of the points that belong to each label are close to the original means.

<!-- grade_id: cell-09a1705f91d1d892, points: 10, grade: true, solution: false -->
```ruby
def test_11_3()
  dataset = create_cluster_dataset()
  examples = dataset["data"]
  clusters = examples.group_by {|x| x["label"]}
  means = [[-4,4], [-4,-4], [4,4], [4,-4], [0,0]]
  n = 5.times do |i|
    assert_not_nil clusters[i]
    assert_equal 100, clusters[i].size
    m1 = clusters[i].inject(0.0) {|u,r| u += r["features"]["x1"]} / 100.0
    m2 = clusters[i].inject(0.0) {|u,r| u += r["features"]["x2"]} / 100.0
    assert_in_delta m1, means[i][0], 0.5
    assert_in_delta m2, means[i][1], 0.5    
  end
  assert_equal 5, n
end
test_11_3()
```

## Question 1.2 (25 points)

Generate $k = 5$ random means. Each selected mean shoulnd be within the range of the min and max of each feature value.

<!-- grade_id: cell-6a934e257380e926, points: 0, grade: false, solution: true -->
```ruby
def init_cluster examples, k
  means = Hash.new {|h,k| h[k] = Hash.new {|h,k| h[k] = 0.0}}
  # BEGIN YOUR CODE
  j = 0
  while j < k do
                clusters = examples.group_by {|x| x["cluster"]}
                size = clusters.size
                a = clusters[j % size]
                n = clusters[j % size].length()

                x1_list = []
                x2_list = []
                i = 0
                while  i < n  do
                                x1_list << a[i]["features"]["x1"]
                                x2_list << a[i]["features"]["x2"]
                                i +=1
                end
                u1 = rand(x1_list.min...x1_list.max)
                u2 = rand(x2_list.min...x2_list.max)
                means[j] = {"x1"=> u1, "x2"=> u2}
                j +=1
  end
  
  
  
  #END YOUR CODE
  means
end
```

Output:
```
:init_cluster
```

<!-- grade_id: cell-52ceae8580209a6d, points: 25, grade: true, solution: false -->
```ruby
def test_12_1()
  dataset = create_cluster_dataset()  
  examples = dataset["data"]
  means = init_cluster examples, 5
  assert_equal 5, means.size
  assert_equal 2, means.first.size
  5.times {|i| assert means.has_key? i}
  assert_true(means.values.all? {|m| m["x1"].abs < 10}, "Expected mean of x1 in [-10, 10]")
  assert_true(means.values.all? {|m| m["x2"].abs < 10}, "Expected mean of x2 in [-10, 10]")
end

test_12_1()
```

```ruby

```

## Question 2.1: (10 points)
We will now implement the k-means algorithm to recover the clusters above. Return the means and plot the obtained clustering. Compare your discovered means to the known means.

Starting from a randomly initialized set of clusters defined by a set of means, we will iteratively refine our estimate of the clustering. At each step, we update the means, assign clusters, and update means from the assigned clustering. We denote candidate cluster membership by the binary vector $z_{i,j} = 1$ if example $i$ belongs to cluster $j$. Use an ```Array``` for a row of z.

Start by setting the $z$ vector by assigning finding the mean closest to each point in the datset. Set the candidate cluster in the ```cluster``` field in each row. 

```ruby
def Closest_index means,x1,x2   ###find which cluster is closest 
    num = means.length()
    distance = []
    i = 0
    while  i < num  do
            distance << ((means[i]["x1"]-x1)**2 + (means[i]["x2"]-x2)**2)**0.5
            i +=1        
    end
    index = distance.each_with_index.min[1]
    return index
end 
```

Output:
```
:Closest_index
```

```ruby
def Get_z examples,means
  num2 = examples.length()
  l = 0
  z = Array.new
  while l <  num2  do
    x1 = examples[l]["features"]["x1"]
    x2 = examples[l]["features"]["x2"]
    index = Closest_index means,x1,x2
    a = Array.new(means.size) {|i| 0.0}
    a[index] = 1.0
    examples[l]["cluster"] = index
    l +=1
    z << a
  end 
  return  z
end
```

Output:
```
:Get_z
```

```ruby
def update_mean z,examples  
            means = Hash.new {|h,k| h[k] = Hash.new {|h,k| h[k] = 0.0}}
            indices = Array.new(means.size) {|i| i}
            
            num1 = z[0].length
            num2 = examples.length()

            m = 0
            while m < num1 do
                      sum1 = 0
                      sum2 = 0
                      sum3 = 0
                      l = 0
                      while l <  num2  do
                        a = (z[l][m]) * (examples[l]["features"]["x1"])
                        b = (z[l][m]) * (examples[l]["features"]["x2"])
                        c = z[l][m]
                        sum1 +=a
                        sum2 +=b
                        sum3 +=c
                        l += 1

                      end

                      u1 = sum1/sum3
                      u2 = sum2/sum3

                      means[m] = {"x1"=> u1, "x2"=> u2}
                      m +=1
            end
  return means
end
```

Output:
```
:update_mean
```

<!-- grade_id: cell-8703280af8c50289, points: 0, grade: false, solution: true -->
```ruby
def assign_cluster examples, means
  indices = Array.new(means.size) {|i| i}
  z = Array.new
  # BEGIN YOUR CODE
  z0 = Get_z examples,means 
  means = update_mean(z0,examples) 
  while z0!=  (Get_z examples,means) do
    z0 = Get_z examples,means
    means = update_mean(z0,examples) 
  end  
  z = Get_z examples,means
  
  
  #END YOUR CODE
  
  return z
end

```

Output:
```
:assign_cluster
```

<!-- grade_id: cell-534b63511cbeda9d, points: 10, grade: true, solution: false -->
```ruby
def test_21_1()
  dataset = create_cluster_dataset()  
  examples = dataset["data"]
  means = init_cluster examples, 5
  z = assign_cluster(examples, means)
  assert_equal 500, z.size
  assert_true(z.all? {|zi| zi.size == 5}, "Must set a value for each cluster")
  assert_true(z.all? {|zi| g = zi.group_by {|zij| zij}; g[0.0].size == 4 and g[1.0].size == 1}, 
    "Must set only one cluster to 1.0 and all others to 0.0 (not an integer)")
  
  plot_clusters(examples)
end
test_21_1()
```

Output:
```
#<Nyaplot::Plot:0x0000555f56e4f830 @properties={:diagrams=>[#<Nyaplot::Diagram:0x0000555f56854700 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>0, :color=>"rgb(166,206,227)", :tooltip_contents=>[0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]}, :data=>"ccdf9017-a4ec-47b2-b7f5-df525581ff06"}, @xrange=[-6.6091270645562865, -1.3497377699306923], @yrange=[1.552005708160153, 6.074674003211595]>, #<Nyaplot::Diagram:0x0000555f57e4f4b0 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>4, :color=>"rgb(31,120,180)", :tooltip_contents=>[4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4, 4
```

## Question 2.2: (5 points)
Given the $z$ vector and the dataset, calculate the means for each cluster determined by $z$. 


```ruby

```

<!-- grade_id: cell-f651c1c0bc2344d7, points: 0, grade: false, solution: true -->
```ruby
def calculate_means z, examples
  means = Hash.new {|h,k| h[k] = Hash.new {|h,k| h[k] = 0.0}}
  # BEGIN YOUR CODE

  means = Hash.new {|h,k| h[k] = Hash.new {|h,k| h[k] = 0.0}}
  indices = Array.new(means.size) {|i| i}

  num1 = z[0].size
  num2 = examples.length()

  m = 0
  while m < num1 do
            sum1 = 0
            sum2 = 0
            sum3 = 0
            l = 0
            while l <  num2  do
              a = z[l][m]* examples[l]["features"]["x1"]
              b = z[l][m]* examples[l]["features"]["x2"]
              c = z[l][m]
              sum1 +=a
              sum2 +=b
              sum3 +=c
              l += 1

            end

            u1 = sum1/sum3
            u2 = sum2/sum3

            means[m] = {"x1"=> u1, "x2"=> u2}
            m +=1
  end

  #END YOUR CODE
  means
end
```

Output:
```
:calculate_means
```

<!-- grade_id: cell-7fff6a6137b36c9b, points: 5, grade: true, solution: false -->
```ruby
def test_22_1()
  dataset = create_cluster_dataset()  
  examples = dataset["data"]
  means1 = init_cluster examples, 5
  z = assign_cluster(examples, means1)

  assert_equal 500, examples.size
  assert_equal 500, z.size

  means2 = calculate_means(z, examples)
  assert_equal 5, means2.size
  5.times {|i| assert means2.has_key? i}
  
  5.times do |i|
    assert_true((means2[i]["x1"] - means1[i]["x1"]).abs > 1e-5, "Expected means to change")
    assert_true((means2[i]["x2"] - means1[i]["x2"]).abs > 1e-5, "Expected means to change")
  end  

  assert_true(means2.values.all? {|fm| fm["x1"].abs < 10}, "Expected mean of x1 in [-10, 10]")
  assert_true(means2.values.all? {|fm| fm["x2"].abs < 10}, "Expected mean of x2 in [-10, 10]")
end

test_22_1()
```

## Question 2.3: (5 points)
As we iteratively refine our means for the clusters, we need to terminate after some criterion. Stop when the cluster distances are less than $\tau = 0.001$ and plot the convergence. First, let's calculate the distance between the set of $k$ means, using squared distance. Use the following formula for the termination criteria:

#  $\frac{1}{k} \sum_{k} \left\lVert \mu^{(k)}_t - \mu^{(k)}_{t - 1} \right\rVert^2 \le \tau$


<!-- grade_id: cell-b7105f7e91a2eba1, points: 0, grade: false, solution: true -->
```ruby
def cluster_dist(m0, m1)
  dist = 0.0
  # BEGIN YOUR CODE
  num = m0.keys.length()
  i = 0
  sum = 0 
  while i < num do
    x1 = m0[i].keys[0]
    y1 = m0[i].keys[1]
    x2 = m1[i].keys[0]
    y2 = m1[i].keys[1]
    
    sum += (m0[i][x1]-m1[i][x2])**2 +  (m0[i][y1]-m1[i][y2])**2
    i +=1
  end 
  dist = sum/num
  
  
  
  
  #END YOUR CODE
  dist
end
```

Output:
```
:cluster_dist
```

```ruby

```

<!-- grade_id: cell-abe63804c50cf21b, points: 5, grade: true, solution: false -->
```ruby
def test_23_1()
  m0 = {0 => {"x" => 1.0, "y" => 2.0}, 1 => {"x" => 2.0, "y" => 3.0}}
  m1 = {0 => {"x" => 2.0, "y" => 2.0}, 1 => {"x" => 3.0, "y" => 4.0}}

  assert_equal 0.0, cluster_dist(m0, m0)
  assert_equal 0.0, cluster_dist(m1, m1)
  assert_equal 1.5, cluster_dist(m0, m1)
end

test_23_1()
```

## Question 2.4: (10 Points)

We are now ready to complete the $k$-means clustering algorithm. Given a dataset, the number of clusters desired, $k$, and the termination condition, $\tau$, initialize the means, update the clusters, and refine the estimates until converged. To ensure that the process terminates, stop after 100 iterations. 

When complete, return an array of distances from the current means to the previous means as well as the last vector of means.

<!-- grade_id: cell-92ac307eaadf9b28, points: 0, grade: false, solution: true -->
```ruby
# def k_means(examples,k,tau)
#   dists = []
#   last_means = []
  
#   # BEGIN YOUR CODE
#   last_means << init_cluster(examples,k) 
#   z =  Get_z(examples,last_means[-1])
#   last_means << calculate_means(z,examples)
#   dists << cluster_dist(last_means[-1], last_means[-2])
#   i = 0
#   while (i<100 && dists[-1] > tau) do  
#         z =  Get_z(examples,last_means[-1])
#         last_means << calculate_means(z,examples)
#         dists << cluster_dist(last_means[-1], last_means[-2])
#         i +=1
#   end


#   last_means = last_means[-1]
#   #END YOUR CODE
#   return [dists, last_means]
# end
```

```ruby
def k_means(examples,k,tau)
  dists = []
  last_means = []  
  # BEGIN YOUR CODE
  means = init_cluster(examples,k) 
  z =  Get_z(examples,means)
  last_means = calculate_means(z,examples)
  
  dists << cluster_dist(last_means,means)

  i = 0
  while (i<100 && dists[-1] > tau) do  
        means = last_means
        z =  Get_z(examples,last_means)
        last_means = calculate_means(z,examples)
        dists << cluster_dist(last_means,means)
        
        i +=1
  end


  #END YOUR CODE
  return [dists, last_means]
end
```

Output:
```
:k_means
```

<!-- grade_id: cell-4e1d634d3eee050f, points: 10, grade: true, solution: false -->
```ruby
def test_24_1()
  dataset = create_cluster_dataset()   
  examples = dataset["data"]
  dists, means, z = k_means examples, 3, 0.001
  
  assert_true(dists.size > 3)
  assert_true(dists[1] < dists[0])
  assert_true(dists[-1] < dists[1])
  assert_true(dists.last <= 0.001)
  iters = Array.new(dists.size) {|i| i }
  df = Daru::DataFrame.new({iters: iters, dists: dists})
  df.plot(type: :line, x: :iters, y: :dists) do |plot, diagram|
    plot.x_label "X"
    plot.y_label "Mean Dist"
    diagram.title "Cluster Convergence"
    plot.legend false
  end
end

test_24_1()
```

Output:
```
#<Nyaplot::Plot:0x0000555f56f6a468 @properties={:diagrams=>[#<Nyaplot::Diagram:0x0000555f56f5a040 @properties={:type=>:line, :options=>{:x=>:iters, :y=>:dists, :title=>"Cluster Convergence"}, :data=>"01fb53a5-7233-4176-a06f-2e2d6ddbea6d"}, @xrange=[0, 9], @yrange=[0.0, 4.012287448335479]>], :options=>{:x_label=>"X", :y_label=>"Mean Dist", :legend=>false, :zoom=>true, :width=>700, :xrange=>[0, 9], :yrange=>[0.0, 4.012287448335479]}}>
```

## Question 2.5 (15 points)
Plot the clusters that result for $k = 3, 5, 6, 10, 20$ clusters.

<!-- grade_id: cell-cc692e56ba45c1d7, points: 4, grade: true, solution: false -->
```ruby
def test_25_1(k)
  dataset = create_cluster_dataset()    
  examples = dataset["data"]
  dists, means, z = k_means examples, k, 0.001
  
  assert_equal(k, means.size)
  assert_true(dists.size > 3)
  assert_true(dists[1] < dists[0])
  assert_true(dists[-1] < dists[1])
  assert_true(dists.last <= 0.001)
  iters = Array.new(dists.size) {|i| i }
  plot_clusters(examples)
end

test_25_1(3)
```

Output:
```
#<Nyaplot::Plot:0x0000555f5719daa0 @properties={:diagrams=>[#<Nyaplot::Diagram:0x0000555f56f60918 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>0, :color=>"rgb(141,211,199)", :tooltip_contents=>[0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]}, :data=>"140f3d74-c665-4368-873b-986a185ba3b0"}, @xrange=[-7.613705187048544, -1.0020217680175225], @yrange=[0.07307985306795005, 6.264496807831007]>, #<Nyaplot::Diagram:0x0000555f5694fb78 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>1, :color=>"rgb(255,255,179)", :tooltip_contents=>[1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1,
```

<!-- grade_id: cell-6ed65b7a4cf777eb, points: 4, grade: true, solution: false -->
```ruby
test_25_1(5)
```

Output:
```
#<Nyaplot::Plot:0x0000555f56e8c5c8 @properties={:diagrams=>[#<Nyaplot::Diagram:0x0000555f56810f50 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>0, :color=>"rgb(102,194,165)", :tooltip_contents=>[0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]}, :data=>"94ee38d6-4659-4324-a643-5d60bbc3f6c0"}, @xrange=[-6.641697984216796, -2.006256325113076], @yrange=[0.9103712859120665, 6.01745180954033]>, #<Nyaplot::Diagram:0x0000555f57e34228 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>1, :color=>"rgb(252,141,98)", :tooltip_contents=>[1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 
```

```ruby
test_25_1(2)
```

Output:
```
#<Nyaplot::Plot:0x0000555f57e28f90 @properties={:diagrams=>[#<Nyaplot::Diagram:0x0000555f574499e8 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>0, :color=>"rgb(166,206,227)", :tooltip_contents=>[0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]}, :data=>"576345e3-ff05-4b27-b6ff-5ad732ffa69c"},
```

<!-- grade_id: cell-a28d8170d846f35c, points: 4, grade: true, solution: false -->
```ruby
test_25_1(6)
```

Output:
```
#<Nyaplot::Plot:0x0000555f5783f9d0 @properties={:diagrams=>[#<Nyaplot::Diagram:0x0000555f571793d0 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>0, :color=>"rgb(228,26,28)", :tooltip_contents=>[0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]}, :data=>"8cbe5248-cc89-4551-a9b8-051be8ddf90a"}, @xrange=[-5.810631426293358, -2.487980862844455], @yrange=[1.1379734336652136, 4.14014442067727]>, #<Nyaplot::Diagram:0x0000555f570c9f70 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>5, :color=>"rgb(55,126,184)", :tooltip_contents=>[5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5]}, :data=>"74bdad01-b7ab-4f39-b34b-db61e5c5a454"}, @xrange=[-5.775756004054253, -2.2087269584178832], @yrange=[4.234447036514846, 6.671489637000366]>, #<Nyaplot::Diagram:0x0000555f56d
```

<!-- grade_id: cell-6b0aa8259e513e3a, points: 4, grade: true, solution: false -->
```ruby
test_25_1(10)
```

Output:
```
#<Nyaplot::Plot:0x0000555f56989030 @properties={:diagrams=>[#<Nyaplot::Diagram:0x0000555f5635ada8 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>0, :color=>"rgb(179,226,205)", :tooltip_contents=>[0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]}, :data=>"0ec9b3c8-ec45-4806-9f9a-035b7ef99202"}, @xrange=[-7.588946646318593, -2.1723262501352822], @yrange=[1.8128094121260538, 5.333050704381137]>, #<Nyaplot::Diagram:0x0000555f560d9138 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>5, :color=>"rgb(253,205,172)", :tooltip_contents=>[5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5]}, :data=>"3f55b85d-3a68-4e6f-89aa-8cabbb317570"}, @xrange=[-5.051698618312554, -1.676934370297964], @yrange=[3.132849716659144, 6.993193271552942]>, #<Nyaplot::Diagram:0x0000555f57
```

<!-- grade_id: cell-686a77c7a91acdb9, points: 4, grade: true, solution: false -->
```ruby
test_25_1(20)
```

Output:
```
#<Nyaplot::Plot:0x0000555f56e9c5b8 @properties={:diagrams=>[#<Nyaplot::Diagram:0x0000555f56daef48 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>5, :color=>"rgb(141,211,199)", :tooltip_contents=>[5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5, 5]}, :data=>"029055cd-f8ae-40d7-9480-cb7fa7e30320"}, @xrange=[-6.061315797500262, -3.8201358162266903], @yrange=[4.17401224584038, 6.222000642259539]>, #<Nyaplot::Diagram:0x0000555f56a1a7d8 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>10, :color=>"rgb(255,255,179)", :tooltip_contents=>[10, 10, 10, 10, 10, 10, 10, 10, 10, 10, 10, 10, 10, 10, 10, 10, 10, 10, 10, 10, 10]}, :data=>"f514242f-9cbb-4185-a5fe-12ffdb28fdc2"}, @xrange=[-3.753742644139681, -1.7863500790402727], @yrange=[3.7175432217735125, 5.652022475476463]>, #<Nyaplot::Diagram:0x0000555f568f6028 @properties={:type=>:scatter, :options=>{:x=>:x1, :y=>:x2, :title=>0, :color=>"rgb(190,186,218)", :tooltip_contents=>[
```

```ruby

```