# CS6140 Assignments

**Instructions**
1. In each assignment cell, look for the block:
 ```
  #BEGIN YOUR CODE
  raise NotImplementedError.new()
  #END YOUR CODE
 ```
1. Replace this block with your solution.
1. Test your solution by running the cells following your block (indicated by ##TEST##)
1. Click the "Validate" button above to validate the work.

**Notes**
* You may add other cells and functions as needed
* Keep all code in the same notebook
* In order to receive credit, code must "Validate" on the JupyterHub server

---

```ruby

```

# Assignment 9: Neural Networks

In this assignment we will implement the [Backpropagation](https://en.wikipedia.org/wiki/Backpropagation) algorithm for a neural network using a computational graph. 

Each node in the network has a reference to is ancestor (not successor nodes), as follows:
![Network configuration](assignment-9.png)

We will be implementing the nodes defined above, including their ```forward``` and ```backward``` passes. In the forward pass, the node reads its input and calculates an ```activation``` which it then retains for the next layer. In the backward pass, the partial derivative of output with respect to a variable is calculated in the ```pderv``` function, recursively through the graph. Nodes are connected by setting the respecitve back pointers in the ```connect``` function.

```ruby
require './assignment_lib.rb'

rng = Distribution::Normal.rng(0,1, 293891)

def connect src, dst
  dst.back_net[src.name] = src
end

concentric_circle_data = concentric_dataset()
xor_data = xor_dataset(rng)
two_means_data = generate_synthetic_data(rng)
nil
```

## Question 1: Input

Implement the ```Input``` node which contains a single column of a vector. The value of the variable held by an ```Input``` object is set by calling the ```activation``` accessor directly. In previous assignments we specify a row in a dataset as a hash of features, such as the following:

```ruby
example = {"features" => {"x1" => 1.0, "x2" => 3.0}}
```

Now, the ```Input``` node holds the value of a single column. Each node has a name which represents the variable is knows about. The above statement is equivalent to initializing two ```Input``` nodes once as follows:

```ruby
x1 = Input.new "x1"
x2 = Input.new "x2"
```

Each subsequent example sets the activation value directly, as follows:

```ruby
x1.activation = 1.0
x2.activation = 3.0
```

## Question 1.1 (5 points)

Implement the ```forward``` method which calculates the activation and stores it in a member variable.

<!-- grade_id: cell-803a358c9669c67b, points: 0, grade: false, solution: true -->
```ruby
class Input
  attr_accessor :activation
  attr_reader :back_net
  attr_reader :name
  
  def initialize name    
    @name = name
    @back_net = nil
  end

  def forward
    # BEGIN YOUR CODE
    raise NotImplementedError.new()
    #END YOUR CODE
  end
end
nil
```

<!-- grade_id: cell-c69afda4ea3e718a, points: 5, grade: true, solution: false -->
```ruby
def test_input_forward
  i = Input.new "x1"
  assert_not_nil i
  assert_nil i.activation

  i.activation = 1.0
  assert_equal 1.0, i.activation

  assert_equal 1.0, i.forward

  i.activation = 2.0
  assert_equal 2.0, i.forward
end
test_input_forward()
```

## Question 1.2 (5 points)

Implement the ```pderv``` method which calculates partial derivative of the input with respect to ```name```. That is, this method calculates:

$\begin{align}
\frac{\partial x}{\partial \text{name}}
\end{align}$

where $x$ is an ```Input``` object.

There are two possible derivative values here, when the ```name``` is the variable represented by the input node and when it isn't.

<!-- grade_id: cell-bfdcb2751f1f30fe, points: 0, grade: false, solution: true -->
```ruby
class Input
  def pderv name
    # BEGIN YOUR CODE
    raise NotImplementedError.new()
    #END YOUR CODE
  end
end
```

<!-- grade_id: cell-f15d4fd1c925703c, points: 5, grade: true, solution: false -->
```ruby
def test_input_pderv
  i = Input.new "x"
  i.activation = 7.1
  
  assert_equal 1.0, i.pderv("x")
  assert_equal 0.0, i.pderv("another_variable")
end
test_input_pderv()
```

## Question 2.1 (5 points)

Implement the $L_2$ loss function assuming that there is an ```incoming``` node whose ```activation``` is known. Feedback is provided by explicitly setting ```feedback```.

<!-- grade_id: cell-48e6dd6a3ab48899, points: 0, grade: false, solution: true -->
```ruby
class L2Loss
  attr_reader :activation
  attr_reader :back_net
  attr_reader :name
  attr_accessor :feedback

  def initialize
    @back_net = Hash.new    
  end
  
  def incoming
    @back_net.values.first
  end
  
  def forward
    # BEGIN YOUR CODE
    raise NotImplementedError.new()
    #END YOUR CODE
  end
end
```

<!-- grade_id: cell-b1909157544e7855, points: 5, grade: true, solution: false -->
```ruby
def test_l2loss_forward
  l2loss = L2Loss.new
  i = Input.new "x1"
  connect i, l2loss
  
  i.activation = 7.5
  assert_in_delta 7.5, l2loss.forward, 1e-4
  assert_in_delta 7.5, l2loss.activation, 1e-4
end

test_l2loss_forward()
```

## Question 2.2 (5 points)

Implement the $L_2$ loss fuction based on the ```activation``` and ```feedback```.

<!-- grade_id: cell-510cfe38c3e486f7, points: 0, grade: false, solution: true -->
```ruby
class L2Loss
  def loss
    # BEGIN YOUR CODE
    raise NotImplementedError.new()
    #END YOUR CODE
  end
end
```

<!-- grade_id: cell-e82d2291f4fc0a73, points: 5, grade: true, solution: false -->
```ruby
def test_l2loss_loss
  l2loss = L2Loss.new
  i = Input.new "x1"
  connect i, l2loss
  
  i.activation = 7.5
  l2loss.forward
  l2loss.feedback = 1.0
  assert_in_delta 21.125, l2loss.loss, 1e-4
end
test_l2loss_loss()
```

## Question 2.3 (5 points)

Implement the partial derivative function $L_2$ loss fuction based on the ```activation``` and ```feedback```. 

<!-- grade_id: cell-4032e548c5f47d7e, points: 0, grade: false, solution: true -->
```ruby
class L2Loss  
  def pderv fname    
    # BEGIN YOUR CODE
    raise NotImplementedError.new()
    #END YOUR CODE
  end
end
```

<!-- grade_id: cell-41229b31f1d209b4, points: 5, grade: true, solution: false -->
```ruby
def test_l2loss_pderv
  l2loss = L2Loss.new
  i = Input.new "x1"
  connect i, l2loss
  
  i.activation = 7.5
  l2loss.forward
  l2loss.feedback = 1.0
  assert_in_delta 6.5, l2loss.pderv("x1"), 1e-4
  assert_in_delta 0.0, l2loss.pderv("some_other_variable"), 1e-4
end
test_l2loss_pderv()
```

## Question 2.4 (5 points)

Implement log loss, where the activation is assumed to take the sigmoid of the input. The partial derivative uses the same form as in [Assignment 4](assignment-4/assignment-4.ipynb).


<!-- grade_id: cell-6ce43dd3ed69c280, points: 0, grade: false, solution: true -->
```ruby
class LogLoss
  attr_reader :activation
  attr_reader :back_net
  attr_reader :name
  attr_accessor :feedback

  def initialize
    @back_net = Hash.new
  end
  
  def incoming
    @back_net.values.first
  end

  def forward    
    # BEGIN YOUR CODE
    raise NotImplementedError.new()
    #END YOUR CODE
  end
  
  def pderv fname    
    # BEGIN YOUR CODE
    raise NotImplementedError.new()
    #END YOUR CODE
  end
  
  def loss
    # BEGIN YOUR CODE
    raise NotImplementedError.new()
    #END YOUR CODE
  end
end

```

<!-- grade_id: cell-e246957f8bf03e35, points: 5, grade: true, solution: false -->
```ruby
def test_logloss
  logloss = LogLoss.new
  i = Input.new "x1"
  connect i, logloss
  
  i.activation = 0.77
  assert_in_delta 0.6835208937, logloss.forward, 1e-4
  
  logloss.feedback = 1
  assert_in_delta -0.3164791063, logloss.pderv("x1"), 1e-4  
  assert_in_delta 0.3804980545, logloss.loss, 1e-4
end
test_logloss()
```

## Question 3.1 (5 points)

Implement the Sigmoid activation function. Given the incoming value $x$, the ```forward``` function calculates ```@activation``` as the sigmoid of x, or:

$
\begin{align}
\sigma(x) = \frac{1}{1+ e^{-x}}
\end{align}
$

In the ```pderv``` function, calculate the partial derivative of the activation function and apply the derivative to the ```back_net``` variables. 

<!-- grade_id: cell-e4f0ca293dc53332, points: 0, grade: false, solution: true -->
```ruby
class Sigmoid
  attr_reader :activation
  attr_reader :name
  attr_reader :back_net
  
  def initialize name = "_"
    @name = name
    @back_net = Hash.new
  end
  
  def incoming
    @back_net.values.first
  end
  
  # BEGIN YOUR CODE
  raise NotImplementedError.new()
  #END YOUR CODE
end

```

<!-- grade_id: cell-92522070f935aa59, points: 5, grade: true, solution: false -->
```ruby
def test_sigmoid
  i = Input.new "x1"
  i.activation = 0.0

  sig = Sigmoid.new "s"
  connect i, sig

  s1 = sig.forward
  assert_in_delta 0.5, s1, 1e-4, "sig(0.5)"
  p1 = sig.pderv "x1"
  assert_in_delta 0.25, p1, 1e-4, "pderv(0.5)"

  
  i.activation = 0.775
  s2 = sig.forward
  assert_in_delta 0.6846015003, s2, 1e-4, "sig(0.775)"
  p2 = sig.pderv "x1"
  assert_in_delta 0.2159222861, p2, 1e-4, "pderv(0.775)"
  
  p3 = sig.pderv "some_other_variable"
  assert_in_delta 0.0, p3, 1e-4, "pderv(other)"
end

test_sigmoid()

```

## Question 3.2 (5 points)

Implement the ReLU activation function. Given the incoming value $x$, the ```forward``` function calculates ```@activation``` as a function of x, as follows:

$
\begin{align}
\text{ReLU}(x) = \max \left\{ 0, x \right\}
\end{align}
$

In the ```pderv``` function, calculate the partial derivative of the activation function and apply the derivative to the ```back_net``` variables. 

<!-- grade_id: cell-832dc16bcc2eb5c9, points: 0, grade: false, solution: true -->
```ruby
class ReLU
  attr_reader :activation
  attr_reader :name
  attr_reader :back_net
  
  def initialize name
    @name = name
    @back_net = Hash.new
  end
  
  def incoming
    @back_net.values.first
  end

  # BEGIN YOUR CODE
  raise NotImplementedError.new()
  #END YOUR CODE
end

```

<!-- grade_id: cell-e1145bcacf02e2c0, points: 5, grade: true, solution: false -->
```ruby
def test_relu
  i = Input.new "x1"
  i.activation = 0.5

  relu = ReLU.new "s"
  connect i, relu

  s1 = relu.forward
  assert_in_delta 0.5, s1, 1e-4, "relu(0.5)"
  p1 = relu.pderv "x1"
  assert_in_delta 1.0, p1, 1e-4, "pderv(0.5)"

  
  i.activation = -0.775
  s2 = relu.forward
  assert_in_delta 0.0, s2, 1e-4, "relu(-0.775)"
  p2 = relu.pderv "x1"
  assert_in_delta 0.0, p2, 1e-4, "pderv(-0.775)"
  
  p3 = relu.pderv "some_other_variable"
  assert_in_delta 0.0, p3, 1e-4, "pderv(other)"
end

test_relu()

```

## Question 4.1 (5 points)

The ```LinearUnit``` calculates the inner product between the input activations and an internal weight vector. The weight vector is stored as a reference that is common across the whole network. Therefore, each feature name needs to be unique. 

The diagram below shows the network configuration for Logstic Regression. Multiple ```Input``` units are connected to the ```LinearUnit```, which connected to the ```Sigmoid``` activation function. The activation function is then connected to the ```L2Loss``` output node. 

![Network configuration](assignment-9-Page-2.png)

In the ```forward``` function, each input node in the ```back_net``` has an activation. Assume there is a pre-defined weight for the bias.

<!-- grade_id: cell-88b5fa776d4a2a31, points: 0, grade: false, solution: true -->
```ruby
class LinearUnit
  attr_reader :activation
  attr_reader :back_net
  attr_reader :name
  
  def initialize name, weights
    @name = name
    @weights = weights
    @back_net = Hash.new    
  end

  def n fname
    [@name, fname].join(".")
  end
  
  def forward
    @activation = @weights[n("bias")]
    # BEGIN YOUR CODE
    raise NotImplementedError.new()
    #END YOUR CODE
    return @activation
  end
end
```

<!-- grade_id: cell-aaf6342ced97a4df, points: 5, grade: true, solution: false -->
```ruby
def test_linear_forward
  weights = Hash.new {|h,k| h[k] = 0.7}
  linear = LinearUnit.new("f", weights)
  assert_not_nil linear
  assert_true(linear.back_net.empty?)
  
  i = Input.new "x1"
  i.activation = 0.2
  
  connect i, linear
  assert_same i, linear.back_net["x1"]
  
  assert_in_delta 0.2 * 0.7 + 0.7, linear.forward, 1e-4
  assert_in_delta 0.7, weights["f.x1"], 1e-4
  assert_in_delta 0.7, weights["f.bias"], 1e-4
  assert_equal 2, weights.size

end

test_linear_forward()
```

## Question 4.2 (5 points)

The partial derivative function for the ```LinearUnit``` has two cases. In the ```pderv_weights``` below, calculate the partial derivative for weight variables.

<!-- grade_id: cell-c6756771988ae544, points: 0, grade: false, solution: true -->
```ruby
class LinearUnit
  def pderv_weights name
    # BEGIN YOUR CODE
    raise NotImplementedError.new()
    #END YOUR CODE
  end
  
  def pderv name
    pderv_weights name
  end
end
```

<!-- grade_id: cell-852aad19e2dea1d6, points: 5, grade: true, solution: false -->
```ruby
def test_linear_pderv_weights
  weights = Hash.new {|h,k| h[k] = 0.7}
  linear = LinearUnit.new("f", weights)
  assert_not_nil linear

  i = Input.new "x1"
  i.activation = 0.2  
  connect i, linear

  i.forward
  act = linear.forward

  assert_in_delta 0.2, linear.pderv_weights("f.x1"), 1e-4
  assert_in_delta 1.0, linear.pderv_weights("f.bias"), 1e-4
end
test_linear_pderv_weights()
```

## Question 4.3 (5 points)

Considering the dataset below, copy **your** ```StochasticGradientDescent``` and train a logistic regression model. Make two changes to your trainer:

1. Keep learning rate is fixed, i.e., do not try to reduce the learning rate.
1. Divide by the size of the mini-batch in the SGD not in gradient calculation

The ```LogisticRegression``` function should implement these steps just like every other model we used:

1. ```forward```: Calls the forward function for the input, linear unit, activation, and output nodes.
1. ```func```: Calculate the total $L_2$ loss under a sigmoid activation function for all examples in the dataset.
1. ```grad```: Calculate the total gradient vector for all parameters in the function. im


Because the implementation above takes one row at a time, it will call the forward operation again in ```func``` and ```grad```.


Paste **your** implementation of the AUCMetric class here.

```ruby
module Metric
  def apply scores
  end
end

plot_dataset(two_means_data)
```

<!-- grade_id: cell-05718ff13e51d5f7, points: 0, grade: false, solution: true -->
```ruby

class AUCMetric 
  include Metric
  # BEGIN YOUR CODE
  raise NotImplementedError.new()
  #END YOUR CODE
end

class StochasticGradientDescent
  # BEGIN YOUR CODE
  raise NotImplementedError.new()
  #END YOUR CODE
end
```

<!-- grade_id: cell-0c8cf1aca128ce0d, points: 0, grade: false, solution: true -->
```ruby
class LogisticRegression
  def initialize weights
    @x1 = Input.new "x1"
    @x2 = Input.new "x2"
    @inner = LinearUnit.new "wx", weights
    @sig = Sigmoid.new "sig"
    @out = L2Loss.new
    
    connect @x1, @inner
    connect @x2, @inner
    connect @inner, @sig
    connect @sig, @out
  end
  
  def predict examples
    examples.collect {|row| forward row}
  end
  
  def forward row
    # BEGIN YOUR CODE
    raise NotImplementedError.new()
    #END YOUR CODE
  end
  
  def func dataset, weights
    # BEGIN YOUR CODE
    raise NotImplementedError.new()
    #END YOUR CODE
  end
  
  def grad data, weights
    g = Hash.new {|h,k| h[k] = 0.0}

    # BEGIN YOUR CODE
    raise NotImplementedError.new()
    #END YOUR CODE
    g
  end
  
  def adjust weights
  end
end
```

<!-- grade_id: cell-f3d313d85586d5dd, points: 10, grade: true, solution: false -->
```ruby
def test_logistic_regression data
  cumulative_loss = 0.0
  i = 0
  weights = Hash.new {|h,k| h[k] = 0.0}
  obj = LogisticRegression.new weights
  sgd = StochasticGradientDescent.new obj, weights, 0.02
  5.times do |epoch|
    data.each_slice(20) do |batch|    
      sgd.update batch
      cumulative_loss += obj.func(batch, sgd.weights)
      i += 1
      puts cumulative_loss / i if i % 100 == 0
    end
  end
  
  puts weights  
  predictions = obj.predict data
  scores = predictions.collect.with_index {|score, i| [score, data[i]["label"] > 0 ? 1 : 0]}
  auc = AUCMetric.new.apply scores
  puts "AUC: #{auc}"
  
  assert_true(0.5 >= (cumulative_loss / i), "Loss < 0.5")
  assert_true(auc > 0.9, "AUC > 0.9")
  
  return obj
end

trained_lr_model = test_logistic_regression(two_means_data)
nil
```

### Observe
The following plot of the decision boundary should show a higher values for the positive class.

```ruby
plot_decision_boundary(two_means_data, trained_lr_model)
```

## Question 5.1 (5 points)

Neural networks connect multiple linear units together. Testing that this works. No additional code needed here, just another test.

<!-- grade_id: cell-ad31f8bd759ae2e8, points: 5, grade: true, solution: false -->
```ruby
def test_hidden_forward
  weights = {
      "h0.x1" => 1.0, "h0.bias" => 0.0,
      "h1.x1" => 2.0, "h1.bias" => 0.0,
      "f0.h0" => 3.0, "f0.h1" => 4.0, "f0.bias" => 0.0
    }
  h0 = LinearUnit.new("h0", weights)
  h1 = LinearUnit.new("h1", weights)
  f0 = LinearUnit.new("f0", weights)
  
  i = Input.new "x1"  
  i.activation = 0.2
  
  connect i, h0
  connect i, h1
  
  connect h0, f0
  connect h1, f0

  i.forward
  h0.forward
  h1.forward
  f0.forward

  assert_in_delta 0.2, h0.activation, 1e-4, "h0"
  assert_in_delta 0.4, h1.activation, 1e-4, "h1"
  assert_in_delta 0.2 * 3 + 0.4 * 4, f0.activation, 1e-4, "f0"
  
end
test_hidden_forward()
```

## Question 5.2 (5 points)

Implement the backpropagation function for the linear unit when the variable requested is _not_ one of the weights. This function is what makes the backpropagation algorithm work. We will also overwrite the ```pderv``` function here to use the two cases.

<!-- grade_id: cell-daf3ec42522a615b, points: 0, grade: false, solution: true -->
```ruby
class LinearUnit
  def pderv_back name
    # BEGIN YOUR CODE
    raise NotImplementedError.new()
    #END YOUR CODE
  end
  
  
  def pderv name
    if name.start_with?(@name + ".") and @weights.has_key? name
      pderv_weights name
    else
      pderv_back name
    end
  end
end
```

<!-- grade_id: cell-1fffbe32495bef9c, points: 3, grade: true, solution: false -->
```ruby
def test_hidden_backward
  weights = {
      "h0.x1" => 1.5, "h0.bias" => 0.0,
      "h1.h0" => 2.0, "h1.bias" => 0.0,
      "h2.h0" => 3.0, "h2.bias" => 0.0,    
      "h3.h1" => 4.0, "h3.h2" => 5.0, "h3.bias" => 0.0    
  }
  
  h0 = LinearUnit.new("h0", weights)
  h1 = LinearUnit.new("h1", weights)
  h2 = LinearUnit.new("h2", weights)
  h3 = LinearUnit.new("h3", weights)
  
  i = Input.new "x1"  
  i.activation = 0.2
  
  connect i, h0
  
  connect h0, h1
  connect h0, h2
  connect h1, h3
  connect h2, h3

  [i, h0, h1, h2, h3].each {|n| n.forward}

  assert_in_delta 0.2, h0.pderv("h0.x1"), 1e-4, "h0"
  assert_in_delta 0.3, h1.pderv("h1.h0"), 1e-4, "h1"
  assert_in_delta 0.3, h2.pderv("h2.h0"), 1e-4, "h2"
  assert_in_delta 4.0 * (2 * 0.2) + 5 * (3 * 0.2), h3.pderv("h0.x1"), 1e-4, "h3"
  
end
test_hidden_backward()
```

<!-- grade_id: cell-91192be992563435, points: 2, grade: true, solution: false -->
```ruby
def test_91192b
  w1 = {"v.x1" => 6.21285855261965, "v.x2" => 5.89754417523916, "v.bias" => -28.0932043104677}
  w2 = {"t.x1" => -6.47151832422693, "t.x2" => -6.25964370262662, "t.bias" => 21.3658877452187}
  w3 = {"u.z1" => 5.79180358051235, "u.z2" => 5.7240375033176, "u.bias" => -2.83249422973696}
  
  x1 = Input.new "x1"
  x2 = Input.new "x2"

  v = LinearUnit.new("v", w1)
  z1 = Sigmoid.new "z1"
  connect x1, v
  connect x2, v
  connect v, z1

  t = LinearUnit.new("t", w2)
  z2 = Sigmoid.new "z2"
  connect x1, t
  connect x2, t
  connect t, z2

  u = LinearUnit.new("u", w3)
  yhat = LogLoss.new 
  
  connect z1, u
  connect z2, u
  connect u, yhat
  
  x1.activation = 1.5
  x2.activation = 2.0
  
  x1.forward; x2.forward
  puts "v(x) = %.4f" % v.forward
  puts "z1(v) = %.4f" % z1.forward
  puts "t(x) = %.4f" % t.forward
  puts "z2(t) = %.4f" % z2.forward
  puts "u(z) = %.4f" % u.forward
  puts "yhat(u) = %.4f" % yhat.forward
  assert_in_delta 0.2449, yhat.activation, 1e-3
  
  puts "Calculate d[j]/dw11"
  yhat.feedback = 1
  puts "d[j]/dw11 = %.4f" % yhat.pderv("v.x1")
  assert_in_delta -0.0061, yhat.pderv("v.x1"), 1e-3
  puts "d[z1]/dw11 = %.4f" % z1.pderv("v.x1")
  assert_in_delta 0.0014, z1.pderv("v.x1"), 1e-3
  puts "d[v]/dw11 = %.4f"% v.pderv("v.x1")  
  assert_in_delta 1.5000, v.pderv("v.x1"), 1e-3
  
  puts "Calculate d[j]/dw21"
  puts "d[j]/dw21 = %.4f" % yhat.pderv("t.x1")
  assert_in_delta -1.3542, yhat.pderv("t.x1"), 1e-3
  puts "d[z2]/dw21 = %.4f" % z2.pderv("t.x1")
  assert_in_delta 0.3133, z2.pderv("t.x1"), 1e-3
  puts "d[t]/dw21 = %.4f"% t.pderv("t.x1")  
  assert_in_delta 1.5000, t.pderv("t.x1"), 1e-3

  puts "Calculate d[j]/dw31"
  puts "d[j]/dw31 = %.4f" % yhat.pderv("u.z1")
  assert_in_delta -0.0007, yhat.pderv("u.z1"), 1e-3
  puts "d[u]/dw31 = %.4f"% u.pderv("u.z1")  
  assert_in_delta 0.0009, u.pderv("u.z1"), 1e-3

  puts "Calculate d[j]/dwQ = 0"
  assert_in_delta 0.0, yhat.pderv("Q"), 1e-3
end
test_91192b()
```

## Question 5.3 (10 points)

Implement the ```NeuralNetwork``` class following the pattern for the ```LogisticRegression``` above, where there are three main operations: ```forward```, ```func```, and ```grad```. 

Assume that ```layers``` contains all nodes properly connected. Each layer is an array of nodes connected appropriately. Even if there is only one node in the layer, it needs to be in an array.

We will verify that a 1-node neural network produces exactly the same result as the LogisticRegression model above.

<!-- grade_id: cell-9d40194ae5b08293, points: 0, grade: false, solution: true -->
```ruby
class NeuralNetwork
  attr_reader :batch_loss
  def initialize layers
    @inputs = Hash.new
    @layers = layers
    @out = @layers[-1][0]
    @batch_loss = 0.0
  end

  def forward example
    # BEGIN YOUR CODE
    raise NotImplementedError.new()
    #END YOUR CODE
  end
    
  def func examples, weights
    # BEGIN YOUR CODE
    raise NotImplementedError.new()
    #END YOUR CODE
  end
  
  def grad examples, weights
    update_inputs examples
    g = Hash.new {|h,k| h[k] = 0.0}
    # BEGIN YOUR CODE
    raise NotImplementedError.new()
    #END YOUR CODE
    g
  end
  
  def predict data
    data.map do |row|
      forward row
    end 
  end
  
  def update_inputs examples
    examples.flat_map {|r| r["features"].keys}
      .uniq
      .reject {|k| @inputs.has_key? k}
      .each do |k|
        @inputs[k] = Input.new k
        @layers.first.each {|f| connect @inputs[k], f}
      end
    @inputs
  end
  
  def adjust weights
  end
end
```

<!-- grade_id: cell-33bd3b385e6b9c48, points: 10, grade: true, solution: false -->
```ruby
def test_neural_network_1_layer examples
  cumulative_loss = 0.0
  i = 0
  weights = Hash.new {|h,k| h[k] = 0.0}
  layers = [
    [LinearUnit.new("wx", weights)],
    [Sigmoid.new("sig")],
    [L2Loss.new]
  ]
    
  connect layers[0][0], layers[1][0]
  connect layers[1][0], layers[2][0]

  obj = NeuralNetwork.new layers
  sgd = StochasticGradientDescent.new obj, weights, 0.02
  5.times do |epoch|
    examples.each_slice(20) do |batch|    
      sgd.update batch
      cumulative_loss += obj.func(batch, sgd.weights)
      i += 1
      puts cumulative_loss / i if i % 100 == 0
    end
  end
  
  puts weights  
  predictions = obj.predict examples
  scores = predictions.collect.with_index {|score, i| [score, examples[i]["label"] > 0 ? 1 : 0]}
  auc = AUCMetric.new.apply scores
  puts "AUC: #{auc}"
  
  assert_true(0.5 >= (cumulative_loss / i), "Loss < 0.5")
  assert_true(auc > 0.9, "AUC > 0.9")
  
  return obj
end

trained_lr_model2 = test_neural_network_1_layer(two_means_data)
nil
```

## Question 6.1 (10 points)
Using a 2-layer neural network achieve an $\log$ training loss of less than 0.3 on the XOR dataset. You will initialize a ```problem``` Hash with a weights initializer and the layers. You are free to choose the activation and any layers. We will standardize on the ```LogLoss``` for this dataset, which you will need to add. Don't forget to connect your nodes. You are also free to use any initializer you want.

### Practice
You may find the [TensorFlow Neural Network Playgroud](https://playground.tensorflow.org) helpful. It runs in the browser and has all the datasets we are using here.

* [XOR](https://playground.tensorflow.org/#dataset=xor)
* [Concentric Cirles](https://playground.tensorflow.org/#dataset=circle)

```ruby
plot_dataset(xor_data)
```

<!-- grade_id: cell-a4e911311c17091e, points: 0, grade: false, solution: true -->
```ruby
def create_problem_xor
  rng = Distribution::Normal.rng(0,0.2,293891)
  # BEGIN YOUR CODE
  raise NotImplementedError.new()
  #END YOUR CODE
  return problem_xor
end
```

<!-- grade_id: cell-f40f324e078fe152, points: 10, grade: true, solution: false -->
```ruby
def test_neural_network_xor problem, examples
  cumulative_loss = 0.0
  i = 0

  obj = NeuralNetwork.new problem["layers"]
  sgd = StochasticGradientDescent.new obj, problem["weights"], problem["learning_rate"]
  500.times do |epoch|
    examples.each_slice(10) do |batch|    
      sgd.update batch
      cumulative_loss += obj.func(batch, sgd.weights)
      i += 1
      puts [i, cumulative_loss / i].join("\t") if i % 100 == 0
      break if cumulative_loss / i < 0.3
    end
    break if cumulative_loss / i < 0.3
  end
  
  predictions = obj.predict examples
  scores = predictions.collect.with_index {|score, i| [score, examples[i]["label"] > 0 ? 1 : 0]}
  auc = AUCMetric.new.apply scores
  puts "AUC: #{auc}"
 
  assert_true(cumulative_loss / i < 0.3)
  assert_true(auc > 0.9)
  return obj
end

trained_xor_model = test_neural_network_xor(create_problem_xor(), xor_data)
nil
```

```ruby
plot_decision_boundary xor_data, trained_xor_model
```

## Question 6.1 (5 points)
Using a 3-layer neural network achieve an $\log$ training loss of less than 0.2 on the concentric circles dataset. You will initialize a ```problem``` Hash with a weights initializer and the layers. You are free to choose the activation and any layers. We will standardize on the ```LogLoss``` for this dataset, which you will need to add. Don't forget to connect your nodes.

```ruby
plot_dataset concentric_circle_data
```

<!-- grade_id: cell-f960c71875436c12, points: 0, grade: false, solution: true -->
```ruby
def create_problem_concentric
  rng = Distribution::Normal.rng(0,0.2,293891)
  # BEGIN YOUR CODE
  raise NotImplementedError.new()
  #END YOUR CODE
  return problem_concentric
end
```

<!-- grade_id: cell-23a5f41e6aaab0d9, points: 10, grade: true, solution: false -->
```ruby
def test_neural_network_concentric problem, examples
  cumulative_loss = 0.0
  i = 0

  obj = NeuralNetwork.new problem["layers"]
  sgd = StochasticGradientDescent.new obj, problem["weights"], problem["learning_rate"]
  500.times do |epoch|
    examples.each_slice(10) do |batch|    
      sgd.update batch
      cumulative_loss += obj.func(batch, sgd.weights)
      i += 1
      puts [i, cumulative_loss / i].join("\t") if i % 100 == 0
      break if cumulative_loss / i < 0.4
    end
    break if cumulative_loss / i < 0.4
  end
  
  predictions = obj.predict examples
  scores = predictions.collect.with_index {|score, i| [score, examples[i]["label"] > 0 ? 1 : 0]}
  auc = AUCMetric.new.apply scores
  puts "AUC: #{auc}"
 
  assert_true(cumulative_loss / i < 0.4)
  assert_true(auc > 0.9)

  
  return obj
end

trained_concentric_model = test_neural_network_concentric(create_problem_concentric(), concentric_circle_data)
nil
```

```ruby
plot_decision_boundary(concentric_circle_data, trained_concentric_model)
```

```ruby

```